# MATH 210 Assignment 2

### Instructions

* There are 3 problems and 15 total points
* Write your solutions in the cells below
* You may work on these problems with others but you must write your solutions on your own
* **Do not import** any Python packages such as `math` or `numpy` to complete this assignment
* Execute the test cells to verify that your solutions pass
* This notebook does not contain all tests for grading (this means that your solution may not be completely correct even if it passes all tests below)
* Submit this notebook to Canvas before **11:59pm Wednesday October 10**

### Problem 1 (5 points)

Write a function called `solution_A` which takes 5 input parameters `A`, `x0`, `eps`, `max_iter` and  `z` (default value `10e-12`), implements Newton's method with initial guess `x0` and returns an approximation of a solution of the equation

$$
\frac{1 + Ax^3}{1 + x^2} = x^3
$$

The function `solution_A` should implement Newton's method on the function

$$
f(x) = 1 + Ax^3 - x^3(1+x^2)
$$

The function `solution_A` may terminate in 3 different ways:

1. Implement Newton's method until it produces a value `x_n` such that `|f(x_n)| < eps` and then return `x_n`.
2. If `|f'(x_n)| < z` (ie. the absolute value of the derivative is nearly 0) at any point in the implementation, the function prints `"Zero derivative. No solution found."` and returns `None`.
3. If the number of iterations exceeds the maximum number of iterations `max_iter`, the function prints `'Exceeded maximum iterations. No solution found.'` and returns `None`.

<!-- grade_id: solution_A, points: 0, grade: false, solution: true -->
```python
def solution_A(A,x0,eps,max_iter,z = 10e-12):

    xn = x0
    
    for i in range(1,max_iter+1):
        
        f = 1 + A*xn**3 - xn**3 - xn**5
        df = 3*A*xn**2 - 3*xn**2 - 5*xn**4
        
        if abs(f) < eps:
            return xn

        if abs(df) < z:
            print("Zero derivative. No solution found.")
            return None
    
        xn = xn - f/df
    
    print("Exceeded maximum iteratons. No solution found.")
    return None
```

<!-- grade_id: solution_A_values1, points: 1, grade: true, solution: false -->
```python
"Check that solution_A returns the correct datatype when solution is found."
assert type(solution_A(2,1,0.001,25)) == float , "Return value should be a float."
print("Problem 1 Test 1: Success!")
```

Output:
```
Problem 1 Test 1: Success!
```

<!-- grade_id: solution_A_values2, points: 2, grade: true, solution: false -->
```python
"Check that solution_A returns the correct datatype when solution is found."
assert type(solution_A(12,3,10e-10,25)) == float , "Return value should be a float."
print("Problem 1 Test 2: Success!")
```

Output:
```
Problem 1 Test 2: Success!
```

<!-- grade_id: solution_A_deriv, points: 1, grade: true, solution: false -->
```python
"Check that solution_A returns the correct datatype when there's a zero derivaitve."
assert solution_A(1,0,0.0001,2) == None , "Return value should be None."
print("Problem 1 Test 3: Success!")
```

Output:
```
Zero derivative. No solution found.
Problem 1 Test 3: Success!
```

<!-- grade_id: solution_A_max_iter, points: 1, grade: true, solution: false -->
```python
"Check that solution_A returns the correct datatype when exceed maximum iterations."
assert solution_A(2,1,10e-14,1) == None , "Return value should be None."
print("Problem 1 Test 4: Success!")
```

Output:
```
Exceeded maximum iteratons. No solution found.
Problem 1 Test 4: Success!
```

### Problem 2 (5 points)

Represent a polynomial $p(x)=a_0+a_1x+a_2x^2+\cdots+a_nx^n$ as a list of coefficients $[a_0,a_1,\dots,a_n]$. Write a function called `poly_critical` which takes 4 input parameters `p`, `a`, `b` and `N` where `p` is a Python list of numbers representing a polynomial $p(x)$, `a` and `b` are numbers defining an interval $[a,b]$ and `N` is a positive integer.

The function `poly_critical` implements $N$ iterations of the bisection method applied to the equation $p'(x)=0$ and returns an approximation of a critical point $c$ where $p'(c)=0$ for $c \in [a,b]$.

For example, if `p = [1,-1,1,0,0,1]` (which represents $p(x) = 1-x+x^2+x^5$), `a=0`, `b=1` and `N=10` then the function returns `0.4212656546685004` which approximates a solution of $5x^4+2x-1=0$.

The function `poly_critical` may terminate in 4 ways:

1. If `len(p) < 2` (`p` is linear), the function should print `'No critical points.'` and return `None`.
2. If either initial endpoint is a critical point, return the endpoint.
3. If values at both endpoints and midpoint have the same sign at any iteration, the function should print `'Bisection method fails.'` and return `None`.
4. The function implements N iterations of the bisection method successfully and returns the midpoint of the Nth subinterval.

<!-- grade_id: poly_critical, points: 0, grade: false, solution: true -->
```python
def poly_critical(p,a,b,N):
    
    dp = [n*p[n] for n in range(1,len(p))]
    
    if len(p) < 2:
        print("No critical points")
        return None
    
    an = a
    bn = b
    
    def dP(x):
        seq = [dp[i]*x**i for i in range(0,len(dp))]
        sum_seq = sum(seq)
        return sum_seq
        
    if abs(dP(an)) < 10e-12:
        print("Initial endpoint found.")
        return an
    
    if abs(dP(bn)) < 10e-12:
        print("Initial endpoint found.")
        return bn
        
    if dP(an)*dP(bn) > 0:
        print("Bisection method fails.")
        return None
    
    for n in range(1,N+1):
            
        mn = (an+bn)/2
            
        if dP(an) > 0 and dP(bn) > 0 and dP(mn) > 0:
            print("Bisection method fails.")
            return None
            
        if dP(an) < 0 and dP(bn) < 0 and dP(mn) < 0:
            print("Bisection method fails.")
            return None
            
        if dP(mn)*dP(an) < 0:
            an = an
            bn = mn
            
        elif dP(mn)*dP(bn) < 0:
            an = mn
            bn = bn
    
    return mn
```

<!-- grade_id: poly_critical_values1, points: 1, grade: true, solution: false -->
```python
"Check that poly_critical returns the correct datatype and value."
assert type(poly_critical([1,-1,1,0,0,1],0,1,100)) == float , "Return value should be a float."
print("Problem 2 Test 1: Success!")
```

Output:
```
Problem 2 Test 1: Success!
```

<!-- grade_id: poly_critical_values2, points: 2, grade: true, solution: false -->
```python
"Check that poly_critical returns the correct datatype and value."
assert type(poly_critical([1,-2,3,4],0,1,20)) == float , "Return value should be a float."
print("Problem 2 Test 2: Success!")
```

Output:
```
Problem 2 Test 2: Success!
```

<!-- grade_id: poly_critical_endpoint, points: 1, grade: true, solution: false -->
```python
"Check that poly_critical returns the endpoint if its a critical point."
assert poly_critical([1,2,-1],0,1,2) == 1.0 , "Return value should be 1.0."
print("Problem 2 Test 3: Success!")
```

Output:
```
Initial endpoint found.
Problem 2 Test 3: Success!
```

<!-- grade_id: poly_critical_none, points: 1, grade: true, solution: false -->
```python
"Check that poly_critical returns None if polynomial is linear."
assert poly_critical([-21,18],-1,1,2) == None , "Return value should be None."
print("Problem 2 Test 4: Success!")
```

Output:
```
Bisection method fails.
Problem 2 Test 4: Success!
```

### Problem 3 (5 points)

Given a finite sequence of positive integers $[a_0,a_1,\dots,a_n]$ (of length $n+1$), define a new finite sequence $[b_0,b_1,\dots, b_n]$ by the recursive formula

$$
b_n = a_{n} + \frac{1}{b_{n-1}}
$$

For example:

\begin{align*}
b_0 &= a_0 \\
b_1 &= a_{1} + \frac{1}{b_0} = a_{1} + \frac{1}{a_0} \\
b_2 &= a_{2} + \frac{1}{b_1} = a_2 + \frac{1}{a_{1} + \frac{1}{a_0}}\\
b_3 &= a_{3} + \frac{1}{b_2} = a_3 + \frac{1}{a_2 + \frac{1}{a_{1} + \frac{1}{a_0}}} \\
& \vdots \\
b_n &= a_{n} + \frac{1}{b_{n-1}} = a_n + \frac{1}{\ddots + \frac{1}{a_0}} 
\end{align*}

Write a function called `sequence_to_fraction` which takes one input parameter `a` (a Python list of positive integers $[a_0,a_1,\dots,a_n]$) and returns the last number $b_n$ in the sequence defined above

$$
b_n = a_n + \frac{1}{a_{n-1} + \frac{1}{\ddots + \frac{1}{a_0}}}
$$


<!-- grade_id: sequence_to_fraction, points: 0, grade: false, solution: true -->
```python
def sequence_to_fraction(a):
    
    if len(a) == 1:
        return a[0]
    
    return a[-1] + 1/sequence_to_fraction(a[:-1])
```

<!-- grade_id: sequence_to_fraction_values1, points: 1, grade: true, solution: false -->
```python
"Check that sequence_to_values returns the correct values."
assert abs(sequence_to_fraction([1,1,1]) - 1.5) < 10e-12
print("Problem 3 Test 1: Success!")
```

Output:
```
Problem 3 Test 1: Success!
```

<!-- grade_id: sequence_to_fraction_values2, points: 2, grade: true, solution: false -->
```python
"Check that sequence_to_values returns the correct values."
assert type(sequence_to_fraction([1,1,1,1,1,1,1,1,1,1,1])) == float
print("Problem 3 Test 2: Success!")
```

Output:
```
Problem 3 Test 2: Success!
```

<!-- grade_id: sequence_to_fraction_values3, points: 2, grade: true, solution: false -->
```python
"Check that sequence_to_values returns the correct values."
assert type(sequence_to_fraction([6,1,1,4,1,1,2,1,2])) == float
print("Problem 3 Test 3: Success!")
```

Output:
```
Problem 3 Test 3: Success!
```

```python

```