# MATH 210 Assignment 4

### Instructions

* There are 5 problems and 25 total points
* Write your solutions in the cells below
* You may work on these problems with others but you must write your solutions on your own
* Use NumPy, SciPy and Matplotlib
* Execute the test cells to verify that your solutions pass
* This notebook does not contain all tests for grading (this means that your solution may not be completely correct even if it passes all tests below)
* Submit this notebook to Canvas before **11:59pm Friday November 16**

```python
import numpy as np
import matplotlib.pyplot as plt
import scipy.linalg as la
%matplotlib inline
```

### Problem 1 (5 points)

Write a function called `max_eigval` which takes a 2D NumPy array $A$ representing a square matrix and returns the maximum eigenvalue of $(A + A^T)^2$. The return value should be a non-negative real number. If $A$ is not square or if $A$ is not 2-dimensional, the function `max_eigval` prints the statement `"Matrix is not square."` and returns `None`.

<!-- grade_id: max_eigval, points: 0, grade: false, solution: true -->
```python
def max_eigval(A):
    
    if A.ndim != 2:
        print("Matrix is not square.")
        return None
    
    m = A[:,0].size
    n = A[0,:].size
    
    if m == n:
        return max(la.eigvalsh((A + A.transpose())@(A + A.transpose())))
    
    else:
        print("Matrix is not square.")
        return None
    
```

<!-- grade_id: max_eigval_type1, points: 1, grade: true, solution: false -->
```python
"Check max_eigval returns the correct datatype for a square matrix."
A = np.random.randint(0,10,(5,5))
assert type(max_eigval(A)) == np.float64 , "Return value should be a NumPy float."
print("Problem 1 Test 1: Success!")
```

Output:
```
Problem 1 Test 1: Success!
```

<!-- grade_id: max_eigval_type2, points: 1, grade: true, solution: false -->
```python
"Check max_eigval returns None for a non-square matrix."
A = np.random.randint(0,10,(5,3))
assert max_eigval(A) == None , "Return value should be None for a non-square matrix."
print("Problem 1 Test 2: Success!")
```

Output:
```
Matrix is not square.
Problem 1 Test 2: Success!
```

<!-- grade_id: max_eigval_type3, points: 1, grade: true, solution: false -->
```python
"Check max_eigval returns None for a non-square matrix."
A = np.random.randint(0,10,5)
assert max_eigval(A) == None , "Return value should be None for a non-square matrix."
print("Problem 1 Test 3: Success!")
```

Output:
```
Matrix is not square.
Problem 1 Test 3: Success!
```

<!-- grade_id: max_eigval_val, points: 2, grade: true, solution: false -->
```python
"Check max_eigval returns correct values."
A = np.eye(5)
assert max_eigval(A) == 4.0 , "Return value should be 4.0 for the identity matrix."
print("Problem 1 Test 4: Success!")
```

Output:
```
Problem 1 Test 4: Success!
```

### Problem 2 (5 points)

Write a function called `eigen_X` which takes input parameters $x$ and $y$, both are 1D NumPy arrays of the same length, and returns the eigenvector of $X^TX$ corresponding to the maximum eigenvalue of $X^TX$ where the matrix $X$ is constructed by the process:

* Compute the average value $\mu$ of the entries of $x$
* Compute the average value $\nu$ of the entries of $y$.
* Create new arrays $\bar{x}$ and $\bar{y}$ such that the entries of $\bar{x}$ are $\bar{x}_k = x_k - \mu$ and also $\bar{y}_k = y_k - \nu$. The result is that the average value of the entries of both $\bar{x}$ and $\bar{y}$ are zero.
* Create a matrix $X = [ \ \bar{x} \  \bar{y} \ ]$. This means $\bar{x}$ is the column at index 0 and $\bar{y}$ is the column at index 1.
* Return the eigenvector $v$ of $X^TX$ corresponding to the maximum eigenvalue of $X^T X$ such that:
    * $v$ is 1-dimensional NumPy array
    * $v$ is normalized $||v||=1$ and $v = [v_0,v_1]$ with $v_0 > 0$ (or $v_1\geq0$ if $v_0=0$)

<!-- grade_id: eigen_X, points: 0, grade: false, solution: true -->
```python
def eigen_X(x,y):
    xa = np.mean(x)
    ya = np.mean(y)
    x_bar = x - xa
    y_bar = y - ya
    X = np.column_stack([x_bar,y_bar])
    A = X.transpose()@X
    w, v = la.eigh(A)
    largest_eigenvector = v[:, np.argmax(w)]
    return largest_eigenvector
    
```

<!-- grade_id: eigen_X_type, points: 1, grade: true, solution: false -->
```python
"Check that eigen_X returns the correct datatype and size."
x = np.random.rand(5)
y = np.random.rand(5)
assert type(eigen_X(x,y)) == np.ndarray , "Return value should be a NumPy array."
assert eigen_X(x,y).shape == (2,) , "Return value should be a 1D NumPy array of length 2."
print("Problem 2 Test 1: Success!")
```

Output:
```
Problem 2 Test 1: Success!
```

<!-- grade_id: eigen_X_value1, points: 2, grade: true, solution: false -->
```python
"Check that eigen_X returns the correct values."
x = np.array([1,-1])
y = np.array([2,0])
assert la.norm(eigen_X(x,y) - (1/2**0.5)*np.array([1,1])) < 10e-10 , "Return value should be a [1/2**0.5,1/2**0.5]."
print("Problem 2 Test 2: Success!")
```

Output:
```
Problem 2 Test 2: Success!
```

<!-- grade_id: eigen_X_value2, points: 2, grade: true, solution: false -->
```python
"Check that eigen_X returns the correct values."
x = np.array([1,-1])
y = np.array([2,0])
assert la.norm(eigen_X(x,y) - (1/2**0.5)*np.array([1,1])) < 10e-10 , "Return value should be a [1/2**0.5,1/2**0.5]."
print("Problem 2 Test 3: Success!")
```

Output:
```
Problem 2 Test 3: Success!
```

### Problem 3 (5 points)


Let $n$ be a positive integer and let $h$ be a small positive number. Let $D_n$ be the square matrix of size $n$ such that

\begin{align*}
D_{n; 0,0} &= -\frac{1}{h} \\
D_{n; 0,1} &= \frac{1}{h} \\
D_{n; i,i+1} &= \frac{1}{2h} \text{ for } i=1,\dots,n-2 \\
D_{n; i+1,i} &= -\frac{1}{2h} \text{ for } i=0,\dots,n-3 \\
D_{n; n-1,n-2} &= -\frac{1}{h} \\
D_{n; n-1,n-1} &= \frac{1}{h} \\
\end{align*}

The subscripts $D_{n; i,j}$ indicate the entry at row index $i$ and column index $j$ of the matrix $D_n$. For example, when $n=4$ we have

$$
D_4 = \frac{1}{2h} \left[ \begin{array}{rrrr}
-2 & 2 & 0 & 0 \\
-1 & 0 & 1 & 0 \\
0 & -1 & 0 & 1 \\
0 & 0 & -2 & 2 \\
\end{array} \right]
$$

Write a function called `solve_D` which takes input parameters $n$, $h$ and $b$ (a 2D NumPy array of size $n$ by 1) and returns a solution of the system of equations

$$
\left( D_n^2 - I_n \right) x = b
$$

where $I_n$ is the identity matrix of size $n$. The return value $x$ is a 2D NumPy array of size $n$ by 1.

<!-- grade_id: solve_D, points: 0, grade: false, solution: true -->
```python
def solve_D(n,h,b):
    
    I = np.eye(n)
    D = np.zeros((n,n))
    D[0,0] = -1/h
    D[0,1] = 1/h
    D[n-1,n-2] = -1/h
    D[n-1,n-1] = 1/h
    
    for i in range(1,n-1):
        D[i,i+1] = 1/(2*h)
        
    for j in range(0,n-2):
        D[j+1,j] = -1/(2*h)
    
    A = D@D - I
    
    answer = la.solve(A,b)
    
    return answer
```

<!-- grade_id: solve_D_type, points: 2, grade: true, solution: false -->
```python
"Check that solve_D returns the correct type."
assert type(solve_D(3,1,np.array([0,1,2]).reshape(3,1))) == np.ndarray, "Return type should be a NumPy array."
assert solve_D(3,1,np.array([0,1,2]).reshape(3,1)).shape == (3,1), "Return type should be a 2D NumPy array of size (3,1) in this case."
print("Problem 3 Test 1: Success!")
```

Output:
```
Problem 3 Test 1: Success!
```

<!-- grade_id: solve_D_values, points: 3, grade: true, solution: false -->
```python
"Check that solve_D returns the correct values."
assert la.norm(solve_D(3,1,np.array([1,-2,4]).reshape(3,1)) - np.array([-5.5,-2.5,-8.5]).reshape(3,1)) < 10e-12
print("Problem 3 Test 2: Success!")
```

Output:
```
Problem 3 Test 2: Success!
```

##  Problem 4 (5 points)

Let $H_n$ be the square matrix of size $n$ such that
$$
H_{n; i,j} = \frac{1}{i + j + 2} \ \  0 \leq i,j < n
$$

The subscripts $H_{n; i,j}$ indicate the entry at row index $i$ and column index $j$ of the matrix $H_n$. For example, when $n=4$ we have

$$
H_4 = \left[ \begin{array}{rrrr}
1/2 & 1/3 & 1/4 & 1/5 \\
1/3 & 1/4 & 1/5 & 1/6 \\
1/4 & 1/5 & 1/6 & 1/7 \\
1/5 & 1/6 & 1/7 & 1/8
\end{array} \right]
$$

Write a function called `solve_H` which takes input parameter $n$ and $b$ (a 2D NumPy array of size $n$ by 1), and returns a solution of the system of equations

$$
\left( H_n^2 - I_n \right) x = b
$$

where $I_n$ is the identity matrix of size $n$. The return value $x$ is a 2D NumPy array of size $n$ by 1.

<!-- grade_id: solve_H, points: 0, grade: false, solution: true -->
```python
def solve_H(n,b):
    
    I = np.eye(n)
    H = np.ones((n,n))
    
    for j in range(0,n):
        for i in range(0,n):
            
            H[i,j] = 1/(i+j+2)
            
    A = H@H - I
    
    answer = la.solve(A,b)
    
    return answer
```

<!-- grade_id: solve_H_type, points: 2, grade: true, solution: false -->
```python
"Check that solve_H returns the correct type."
assert type(solve_H(10,np.arange(10)).reshape(10,1)) == np.ndarray, "Return type should be a NumPy array."
assert solve_H(10,np.arange(10)).reshape(10,1).shape == (10,1), "Return type should be a 2D NumPy array of size 10 by 1 in this case."
print("Problem 4 Test 1: Success!")
```

Output:
```
Problem 4 Test 1: Success!
```

<!-- grade_id: solve_H_values, points: 3, grade: true, solution: false -->
```python
"Check that solve_H returns the correct type."
x = np.array([-3.27992737,-3.33209729,-3.81807645]).reshape(3,1)
assert la.norm(solve_H(3,np.arange(3).reshape(3,1)) - x) < 10e-8, "Return type should be a NumPy array."
print("Problem 4 Test 2: Success!")
```

Output:
```
Problem 4 Test 2: Success!
```

##  Problem 5 (5 points)

Write a function called `p_a` which takes 2 input parameters `y` and `a` where `y` is a 1D NumPy array of length $N+1$ and `a` is a number. The function returns the value $p(a)$ where $p(x)$ is the unique polynomial of degree $N$ interpolating the points $(i,y_i)$, $i=0,\dots,N$.

<!-- grade_id: p_a, points: 0, grade: false, solution: true -->
```python
def p_a(y,a):
    
    n = len(y) - 1
    x = np.arange(0,n+1,1)
    A = np.vander(x)
    xn = la.solve(A,y)
    answer = sum([xn[k]*(a**(n-k)) for k in range(0,n+1)])
    return answer
```

<!-- grade_id: p_a_values1, points: 1, grade: true, solution: false -->
```python
"Check that solve_D returns the correct values."
assert np.abs(p_a(np.array([1,0,1]),3) - 4) < 1e-10
print("Problem 5 Test 1: Success!")
```

Output:
```
Problem 5 Test 1: Success!
```

<!-- grade_id: p_a_values2, points: 2, grade: true, solution: false -->
```python
"Check that solve_D returns the correct values."
assert np.abs(p_a(np.array([1,0,1]),3) - 4) < 1e-10
print("Problem 5 Test 2: Success!")
```

Output:
```
Problem 5 Test 2: Success!
```

<!-- grade_id: p_a_values3, points: 2, grade: true, solution: false -->
```python
"Check that solve_D returns the correct values."
y = np.array([-2, -4,  2,  3,  2,  2])
assert np.abs(p_a(y,2.5) - 3.125) < 1e-10
print("Problem 5 Test 3: Success!")
```

Output:
```
Problem 5 Test 3: Success!
```