# MATH 210 Assignment 1

### Instructions

* There are 4 problems and 18 total points
* Write your solutions in the cells below
* You may work on these problems with others but you must write your solutions on your own
* **Do not import** any Python packages such as `math` or `numpy` to complete this assignment
* Execute the test cells to verify that your solutions pass
* This notebook does not contain all tests for grading (this means that your solution may not be completely correct even if it passes all tests below)
* Submit this notebook to Canvas before **11:59pm Friday September 21**

### Problem 1 (3 points)

Consider the beautiful infinite series identity

$$
\frac{2 \pi}{e^{\pi x} - e^{-\pi x}} = x \sum_{k=-\infty}^{\infty} \frac{(-1)^k}{x^2 + k^2}
$$

Write a function called `k_sum` which takes two input parameters `x` and `N` and returns the sum

$$
x \sum_{k=-N}^{N} \frac{(-1)^k}{x^2 + k^2}
$$

<!-- grade_id: k_sum, points: 0, grade: false, solution: true -->
```python
def k_sum(x,N) :
    terms = [(x*((-1)**k))/(x**2 + k**2) for k in range(-N,N+1)]
    partial_sum = sum(terms)
    return partial_sum
    
```

<!-- grade_id: k_sum_type, points: 1, grade: true, solution: false -->
```python
"Check that k_sum returns the correct datatype."
assert type(k_sum(1,1)) == float , "Return value should be a float."
print("Problem 1 Test 1: Success!")
```

Output:
```
Problem 1 Test 1: Success!
```

<!-- grade_id: k_sum_values, points: 2, grade: true, solution: false -->
```python
"Check that k_sum returns the correct values."
assert k_sum(1,0) == 1.0 , "Value k_sum(1,0) should be equal to 1.0."
print("Problem 1 Test 2: Success!")
```

Output:
```
Problem 1 Test 2: Success!
```

### Problem 2 (3 points)

Write a function called `log_taylor` which takes one input parameter `N` and returns a Python list of length `N` representing the sequence $\{ a_1,a_2,\dots,a_N\}$ of coefficients of the Taylor polynomial $T_N(x)$ of degree $N$ of $\ln(1 + x)$

$$
T_N(x) = \sum_{k=1}^N \frac{(-1)^{k+1}}{k} x^k
$$

In other words, the function returns the list $[a_1,a_2,\dots,a_N]$ where

$$
a_k = \frac{(-1)^{k+1}}{k}
$$

<!-- grade_id: log_taylor, points: 0, grade: false, solution: true -->
```python
def log_taylor(N) :
    numbers = [((-1)**(k+1))/k for k in range (1,N+1)]
    numbers_list = list(numbers)
    return numbers_list
```

<!-- grade_id: log_taylor_type, points: 1, grade: true, solution: false -->
```python
"Check that log_taylor returns the correct datatype."
assert type(log_taylor(3)) == list , "Return value should be a list."
assert len(log_taylor(3)) == 3 , "Return value should be a list of length 3 when N=3."
print("Problem 2 Test 1: Success!")
```

Output:
```
Problem 2 Test 1: Success!
```

<!-- grade_id: log_taylor_values, points: 2, grade: true, solution: false -->
```python
"Check that log_taylor returns the correct values."
assert log_taylor(1) == [1.0] , "Value log_taylor(1) should be [1.0]."
print("Problem 2 Test 2: Success!")
```

Output:
```
Problem 2 Test 2: Success!
```

### Problem 3 (4 points)

Represent a polynomial $p(x)=a_0+a_1x+a_2x^2+\cdots+a_nx^n$ as a list of coefficients $[a_0,a_1,a_2,\dots,a_n]$. Write a function called `poly_eval` which takes two input parameters `p` and `c`, where `p` is a Python list of numbers representing a polynomial $p(x)$ and `c` is a number, and returns the value $p(c)$.

<!-- grade_id: poly_eval, points: 0, grade: false, solution: true -->
```python
def poly_eval(p,c) :
    n = len(p)
    poly_num = [p[i]*(c**i) for i in range(0,n)]
    poly_sum = sum(poly_num)
    return poly_sum
```

<!-- grade_id: poly_eval_values1, points: 1, grade: true, solution: false -->
```python
"Check that poly_eval returns the correct values."
assert poly_eval([1,1,1],1) == 3
print("Problem 3 Test 1: Success!")
```

Output:
```
Problem 3 Test 1: Success!
```

<!-- grade_id: poly_eval_values2, points: 3, grade: true, solution: false -->
```python
"Check that poly_eval returns the correct values."
assert poly_eval([1,0,1],2) == 5
print("Problem 3 Test 2: Success!")
```

Output:
```
Problem 3 Test 2: Success!
```

### Problem 4  (8 points)

Write a function called `statistics` which takes one input parameter `samples`, a list of numbers, and returns a tuple (of length 7) of statistics:

    (mean, standard deviation, minimum, median, maximum, range, total number of samples)

Recall, if `samples` is the list of numbers $[x_0,x_1,...,x_N]$ then

$$
\begin{align}
\text{mean} & : & \mu = \frac{1}{N+1}\sum_{k=0}^N x_k \\
\text{standard deviation} & : & \sqrt{\frac{\sum_{k=0}^N (\mu - x_k)^2}{N + 1}} \\
\text{range} & : & \max\{x_k\} - \min\{x_k\} \\
\end{align}
$$

The [median](https://en.wikipedia.org/wiki/Median) is the middle value if the number of samples is odd, or the median is the average of the two middle values if the number of samples is even. Use the built-in function `sorted` to sort the list of values. 

Remember:
* Do not import any Python packages
* Do not use Python keywords or built-in function names (such as `max`, `min`, and `range`) as variable names

<!-- grade_id: statistics, points: 0, grade: false, solution: true -->
```python
def statistics(samples) :
    n = len(samples)
    samples = sorted(samples)
    stat_mean = (sum([samples[i] for i in range(0,n)]))/n
    stat_stdv = ((sum([(stat_mean - samples[i])**2 for i in range(0,n)]))/(n))**0.5
    stat_minimum = min(samples)
    stat_maximum = max(samples)
    
    if n % 2 == 0 :
        stat_median = (samples[(n-1)/2] + samples[(n+1)/2])/2
        
    if n % 2 == 1 :
        stat_median = samples[(n-1)//2]
            
    stat_range = stat_maximum - stat_minimum
    final_tuple = (stat_mean, stat_stdv, stat_minimum, stat_median, stat_maximum, stat_range, n) 
    return final_tuple
```

<!-- grade_id: statistics_type, points: 1, grade: true, solution: false -->
```python
"Check that statistics returns the correct datatype."
assert type(statistics([1.3,5.2,3.7])) == tuple , "Return value should be a tuple."
assert len(statistics([1.3,5.2,3.7])) == 7 , "Return value should be a tuple of length 7."
print("Problem 4 Test 1: Success!")
```

Output:
```
Problem 4 Test 1: Success!
```

<!-- grade_id: statistics_mean, points: 1, grade: true, solution: false -->
```python
"Check that statistics returns the correct mean value."
epsilon = 10e-14
assert abs(statistics([1.0,2.0,3.0])[0] - 2.0) < epsilon , "Mean value failed."
print("Problem 4 Test 2: Success!")
```

Output:
```
Problem 4 Test 2: Success!
```

<!-- grade_id: statistics_median, points: 2, grade: true, solution: false -->
```python
"Check that statistics returns the correct median value."
assert abs(statistics([1.0,2.0,3.0])[3] - 2.0) < epsilon , "Median value failed."
print("Problem 4 Test 3: Success!")
```

Output:
```
Problem 4 Test 3: Success!
```

<!-- grade_id: statistics_std, points: 2, grade: true, solution: false -->
```python
"Check that statistics returns the correct standard deviation value."
epsilon = 10e-14
assert abs(statistics([1.0,2.0,3.0])[1] - (2/3)**0.5) < epsilon , "Standard deviation value failed."
print("Problem 4 Test 4: Success!")
```

Output:
```
Problem 4 Test 4: Success!
```

<!-- grade_id: statistics_range, points: 2, grade: true, solution: false -->
```python
"Check that statistics returns the correct range values."
epsilon = 10e-14
assert abs(statistics([1.0,2.0,3.0])[2] - 1.0) < epsilon , "Minimum value failed."
assert abs(statistics([1.0,2.0,3.0])[4] - 3.0) < epsilon , "Maximum value failed."
assert abs(statistics([1.0,2.0,3.0])[5] - 2.0) < epsilon , "Range value failed."
assert statistics([1.0,2.0,3.0])[6] == 3 , "Total value failed."
print("Problem 4 Test 5: Success!")
```

Output:
```
Problem 4 Test 5: Success!
```

```python

```

```python

```